<?xml version="1.0"?>
<problem type="lecture" showanswer="attempted" rerandomize="true" title="Generating a dynamic plot with scripts" display_name="Dynamic plot with scripts">
  <text><h2>Example problem with dynamic plot</h2><p>
This problem demonstrates how a plot can be generated dynamically, and used as the basis for a student question.  The plot is generated by code in a &lt;script&gt; stanza.   The code is also used for a customresponse question.
</p><p>The plot also requires a bit of HTML, which is stored in a separate
file (in this case, <tt><a href="/static/html/x_dynamic_plot.html">x_dynamic_plot.html</a></tt>).  Google's <a href="https://google-developers.appspot.com/chart/interactive/docs/gallery/linechart">chart
tools</a> are used to create the plot from data in the script.</p><script type="loncapa/python">
title="Example plot generated by script"
xdat = range(30) 
ydat = [ random.randint(0,50) for x in xdat ]
columns = ['x-value','y-value']
data = repr([columns] + [list(x) for x in zip(xdat,ydat)])

the_ans = ydat[0]

def test_ans(expect,ans):
    return abs(int(ans)-the_ans) &lt; 5


</script>

The script used is:
<blockquote><pre>
&lt;script type="loncapa/python"&gt;&lt;![CDATA[
title="Example plot generated by script"
xdat = range(30) 
ydat = [ random.randint(0,50) for x in xdat ]
columns = ['x-value','y-value']
data = repr([columns] + [list(x) for x in zip(xdat,ydat)])

the_ans = ydat[0]

def test_ans(expect,ans):
    return abs(int(ans)-the_ans) &lt; 5

]]&gt;
&lt;/script&gt;
</pre></blockquote>

<br/>
<p>The plot is generated by this code:</p>
<pre>
    &lt;div id="graphX" title="$&#x200B;title" data="$&#x200B;data" /&gt;
    &lt;iframe src="/static/html/x_dynamic_plot.html" width="920" height="520" /&gt;
</pre>
<br/>

<p>And the question is generated by this code:</p>
<pre>
    &lt;customresponse cfn="test_ans"&gt;
        &lt;textline size="40" correct_answer="$the_ans"/&gt;&lt;br/&gt;
    &lt;/customresponse&gt;
</pre>

<div id="graphX" title="$title" data="$data"/>

<iframe src="/static/html/x_dynamic_plot.html" width="920" height="520"/>

Enter an integer close to the value of y at x=0: <br/>
<customresponse cfn="test_ans"><textline size="40" correct_answer="$the_ans"/><br/></customresponse>

</text>
</problem>
